---
id: bad87fee1348bd9aedf08822
title: Ajuster la marge d'un élément
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJarHW'
forumTopicId: 16654
dashedName: adjust-the-margin-of-an-element
---

# --description--

La "marge" d'un élément contrôle l'espace entre la "bordure" de l'élément et les éléments qui l'entourent.

Ici, nous pouvons voir que la boîte bleue et la boîte rouge sont imbriquées dans la boîte jaune. Notez que la boîte rouge a une "marge" plus grande que la boîte bleue, ce qui la fait paraître plus petite.

Lorsque vous augmentez la marge de la boîte bleue, vous augmentez la distance entre sa bordure et les éléments qui l'entourent.

# --instructions--

Changez la "marge" de la boîte bleue pour qu'elle corresponde à celle de la boîte rouge.

# --hints--

Votre classe `blue-box` devrait donner aux éléments `20px` de `margin`.

```js
assert($('.blue-box').css('margin-top') === '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
